### str

<!-- grade_id: cell-5db49a51603c3906, points: 0, grade: false, solution: true -->
```python
# sample_string = 'Hello'

### BEGIN SOLUTION
sample_string = 'Hello'
### END SOLUTION
```

<!-- grade_id: cell-8ce1c8adb2dc6ec3, points: 1, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_string)) == "<class 'str'>", "type of sample_string is not str"
assert str(len(sample_string)) == "5", "length of sample_string is not correct"
assert str(sample_string.lower()) == "hello", "value of sample_string is not correct"
assert str(sample_string) == "Hello", "correct string value of sample_string but incorrect case of letters"

assert sha1(str(type(sample_string)).encode("utf-8")+b"e492e257e18cd41a").hexdigest() == "ad7da43843e2b509d92231d9d42a07fbb1a5ff1d", "type of sample_string is not str"
assert sha1(str(len(sample_string)).encode("utf-8")+b"e492e257e18cd41a").hexdigest() == "a26b65bd32db60a24b4de335449b1bd3f493c63d", "length of sample_string is not correct"
assert sha1(str(sample_string.lower()).encode("utf-8")+b"e492e257e18cd41a").hexdigest() == "4897f8d87b6eb9b4cfb244551d0333623ad80c51", "value of sample_string is not correct"
assert sha1(str(sample_string).encode("utf-8")+b"e492e257e18cd41a").hexdigest() == "6b65c06b42413592cae5efc401548ee65c65904f", "correct string value of sample_string but incorrect case of letters"

print('Success!')
```

### int

<!-- grade_id: cell-20ef3939e635dabb, points: 0, grade: false, solution: true -->
```python
# sample_int = 35

### BEGIN SOLUTION
sample_int = 35
### END SOLUTION
```

<!-- grade_id: cell-138112d9b5ae4caa, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_int)) == "<class 'int'>", "type of sample_int is not int"
assert str(sample_int) == "35", "value of sample_int is not correct"

assert sha1(str(type(sample_int)).encode("utf-8")+b"634d798e400f4fc5").hexdigest() == "3748ff140e045f3bd69e62aae71ad0604be4ba83", "type of sample_int is not int"
assert sha1(str(sample_int).encode("utf-8")+b"634d798e400f4fc5").hexdigest() == "0497bd5bb6c0b05e52c3a9c5b25c6c156c8f82d6", "value of sample_int is not correct"

print('Success!')
```

### float

<!-- grade_id: cell-0ad2fc0d71a06f59, points: 0, grade: false, solution: true -->
```python
# sample_float = 35.98765345

### BEGIN SOLUTION
sample_float = 35.98765345
### END SOLUTION
```

<!-- grade_id: cell-126da5a3faa851f1, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_float)) == "<class 'float'>", "type of sample_float is not float"
assert str(round(sample_float, 8)) == "35.98765345", "value of sample_float is not correct (rounded to 8 decimal places)"

assert sha1(str(type(sample_float)).encode("utf-8")+b"b59f6419a81c5968").hexdigest() == "9eb13f95967e0e06d97800dda088317ee3932147", "type of sample_float is not float"
assert sha1(str(round(sample_float, 8)).encode("utf-8")+b"b59f6419a81c5968").hexdigest() == "1599c494b0225bb187b51b757041c4096c2fcf21", "value of sample_float is not correct (rounded to 8 decimal places)"

print('Success!')
```

### list

<!-- grade_id: cell-43678e65bae42fd4, points: 0, grade: false, solution: true -->
```python
# sample_list = [1, 'b', "a", 4.5]

### BEGIN SOLUTION
sample_list = [1, 'b', "a", 4.5]
### END SOLUTION
```

<!-- grade_id: cell-c474a81d248cbe37, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_list)) == "<class 'list'>", "type of sample_list is not list"
assert str(len(sample_list)) == "4", "length of sample_list is not correct"
assert str(sorted([str(element) for element in sample_list])) == "['1', '4.5', 'a', 'b']", "values of sample_list are not correct"
assert str(sample_list) == "[1, 'b', 'a', 4.5]", "order of elements of sample_list is not correct"

assert sha1(str(type(sample_list)).encode("utf-8")+b"61f599e0db16d5a8").hexdigest() == "6d05fc6d5068d10d78064cb16eb33875c1b4d0c1", "type of sample_list is not list"
assert sha1(str(len(sample_list)).encode("utf-8")+b"61f599e0db16d5a8").hexdigest() == "ed990778f7fbc23470989eebc133479bd57c67fd", "length of sample_list is not correct"
assert sha1(str(sorted([str(element) for element in sample_list])).encode("utf-8")+b"61f599e0db16d5a8").hexdigest() == "5a5af53b9eb170ba84d62db3eadbe72e4b0c1d92", "values of sample_list are not correct"
assert sha1(str(sample_list).encode("utf-8")+b"61f599e0db16d5a8").hexdigest() == "c95d36c254088dc91b7f6e64c0265523bb407f69", "order of elements of sample_list is not correct"

print('Success!')
```

### set

<!-- grade_id: cell-fe5861008f297022, points: 0, grade: false, solution: true -->
```python
# sample_set = {1, 2, 3, 4, 5}

### BEGIN SOLUTION
sample_set = {1, 2, 3, 4, 5}
### END SOLUTION
```

<!-- grade_id: cell-5e1c667195b47187, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_set)) == "<class 'set'>", "type of sample_set is not set"
assert str(len(sample_set)) == "5", "length of sample_set is not correct"
assert str(sample_set) == "{1, 2, 3, 4, 5}", "value of sample_set is not correct"

assert sha1(str(type(sample_set)).encode("utf-8")+b"10e1b0d0051a49bd").hexdigest() == "9e64588c6438c908f4b4454c75913f2917819e7d", "type of sample_set is not set"
assert sha1(str(len(sample_set)).encode("utf-8")+b"10e1b0d0051a49bd").hexdigest() == "9193928e65af6d6879770a0736275f628e5d03a4", "length of sample_set is not correct"
assert sha1(str(sample_set).encode("utf-8")+b"10e1b0d0051a49bd").hexdigest() == "f84f82f84752254dbca88022884de637e69b3fbe", "value of sample_set is not correct"

print('Success!')
```

### tuple

<!-- grade_id: cell-a971551f5f91de4b, points: 0, grade: false, solution: true -->
```python
# sample_tuple = (23, 78)

### BEGIN SOLUTION
sample_tuple = (23, 78)
### END SOLUTION
```

<!-- grade_id: cell-48cfa5ac37cdd98b, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_tuple)) == "<class 'tuple'>", "type of sample_tuple is not tuple"
assert str(len(sample_tuple)) == "2", "length of sample_tuple is not correct"
assert str(sorted([str(element) for element in sample_tuple])) == "['23', '78']", "values of sample_tuple are not correct"
assert str(sample_tuple) == "(23, 78)", "order of elements of sample_tuple is not correct"

assert sha1(str(type(sample_tuple)).encode("utf-8")+b"7c305a766b45fe54").hexdigest() == "968dfa73b37938a8d368d674c421e66615d1ff3c", "type of sample_tuple is not tuple"
assert sha1(str(len(sample_tuple)).encode("utf-8")+b"7c305a766b45fe54").hexdigest() == "13af45850dae43d012f2566714d19e01a743f9ce", "length of sample_tuple is not correct"
assert sha1(str(sorted([str(element) for element in sample_tuple])).encode("utf-8")+b"7c305a766b45fe54").hexdigest() == "78ef639cee661b27c9acfda3648ea42d30640808", "values of sample_tuple are not correct"
assert sha1(str(sample_tuple).encode("utf-8")+b"7c305a766b45fe54").hexdigest() == "ae0a66649c15f120c0ae29e5a3e910bf9d556afb", "order of elements of sample_tuple is not correct"

print('Success!')
```

### dict

<!-- grade_id: cell-ddabbd9234db36c0, points: 0, grade: false, solution: true -->
```python
# sample_dict = {'a': 1, 'b': 2, "c": 3}

### BEGIN SOLUTION
sample_dict = {'a': 1, 'b': 2, "c": 3}
### END SOLUTION
```

<!-- grade_id: cell-b8b1ad86fbbe578f, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_dict)) == "<class 'dict'>", "type of sample_dict is not dict"
assert str(len(list(sample_dict.keys()))) == "3", "number of keys of sample_dict is not correct"
assert str(sorted(list(sample_dict.keys()))) == "['a', 'b', 'c']", "keys of sample_dict are not correct"
assert str(sorted(list(sample_dict.values()))) == "[1, 2, 3]", "correct keys, but values of sample_dict are not correct"
assert str(sample_dict) == "{'a': 1, 'b': 2, 'c': 3}", "correct keys and values, but incorrect correspondence in keys and values of sample_dict"

assert sha1(str(type(sample_dict)).encode("utf-8")+b"5dc26d9bbd922df6").hexdigest() == "8842b4a36fc4675df9e090b02a4a765999f66714", "type of sample_dict is not dict"
assert sha1(str(len(list(sample_dict.keys()))).encode("utf-8")+b"5dc26d9bbd922df6").hexdigest() == "71586873e3fc3b33499d1c3217b8acc385de2f8c", "number of keys of sample_dict is not correct"
assert sha1(str(sorted(list(sample_dict.keys()))).encode("utf-8")+b"5dc26d9bbd922df6").hexdigest() == "ee4c98b8ea87d25147afd46b041b8bd5e0761e50", "keys of sample_dict are not correct"
assert sha1(str(sorted(list(sample_dict.values()))).encode("utf-8")+b"5dc26d9bbd922df6").hexdigest() == "dcc2156e1c47ddeae9996515cd8250cfbb003e36", "correct keys, but values of sample_dict are not correct"
assert sha1(str(sample_dict).encode("utf-8")+b"5dc26d9bbd922df6").hexdigest() == "aa54f7b7ee244d9a92cd5608ccce0f4740f3985c", "correct keys and values, but incorrect correspondence in keys and values of sample_dict"

print('Success!')
```

### bool

<!-- grade_id: cell-e807929d33f8deb6, points: 0, grade: false, solution: true -->
```python
# sample_bool = True

### BEGIN SOLUTION
sample_bool = True
### END SOLUTION
```

<!-- grade_id: cell-cf2f289017808ac8, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_bool)) == "<class 'bool'>", "type of sample_bool is not bool"
assert str(sample_bool) == "True", "boolean value of sample_bool is not correct"

assert sha1(str(type(sample_bool)).encode("utf-8")+b"b28f82ec52a54219").hexdigest() == "bdacef4b45b422ee595fe265875872fa9826e866", "type of sample_bool is not bool"
assert sha1(str(sample_bool).encode("utf-8")+b"b28f82ec52a54219").hexdigest() == "d188a548ca5e96888925cffd788a4df8f69507ea", "boolean value of sample_bool is not correct"

print('Success!')
```

```python

```