For this problem set, we'll be using the Jupyter notebook:

![](jupyter.png)

---
## Part A (2 points)

Write a function that returns a list of numbers, such that $x_i=i^2$, for $1\leq i \leq n$. Make sure it handles the case where $n<1$ by raising a `ValueError`.

<!-- grade_id: squares, points: 0, grade: false, solution: true -->
```python
def squares(n):
    """Compute the squares of numbers from 1 to n, such that the 
    ith element of the returned list equals i^2.
    
    """
    ### BEGIN SOLUTION
    if n < 1:
        raise ValueError("n must be greater than or equal to 1")
    return [i ** 2 for i in range(1, n + 1)]
    ### END SOLUTION
```

Your function should print `[1, 4, 9, 16, 25, 36, 49, 64, 81, 100]` for $n=10$. Check that it does:

```python
squares(10)
```

<!-- grade_id: correct_squares, points: 1, grade: true, solution: false -->
```python
"""Check that squares returns the correct output for several inputs"""
from hashlib import sha1
assert str(type(squares(1))) == "<class 'list'>", "type of squares(1) is not list"
assert str(len(squares(1))) == "1", "length of squares(1) is not correct"
assert str(sorted([str(element) for element in squares(1)])) == "['1']", "values of squares(1) are not correct"
assert str(squares(1)) == "[1]", "order of elements of squares(1) is not correct"

assert str(type(squares(2))) == "<class 'list'>", "type of squares(2) is not list"
assert str(len(squares(2))) == "2", "length of squares(2) is not correct"
assert str(sorted([str(element) for element in squares(2)])) == "['1', '4']", "values of squares(2) are not correct"
assert str(squares(2)) == "[1, 4]", "order of elements of squares(2) is not correct"

assert str(type(squares(15))) == "<class 'list'>", "type of squares(15) is not list"
assert str(len(squares(15))) == "15", "length of squares(15) is not correct"
assert str(sorted([str(element) for element in squares(15)])) == "['1', '100', '121', '144', '16', '169', '196', '225', '25', '36', '4', '49', '64', '81', '9']", "values of squares(15) are not correct"
assert str(squares(15)) == "[1, 4, 9, 16, 25, 36, 49, 64, 81, 100, 121, 144, 169, 196, 225]", "order of elements of squares(15) is not correct"

assert str(type(squares(10))) == "<class 'list'>", "type of squares(10) is not list"
assert str(len(squares(10))) == "10", "length of squares(10) is not correct"
assert str(sorted([str(element) for element in squares(10)])) == "['1', '100', '16', '25', '36', '4', '49', '64', '81', '9']", "values of squares(10) are not correct"
assert str(squares(10)) == "[1, 4, 9, 16, 25, 36, 49, 64, 81, 100]", "order of elements of squares(10) is not correct"

assert str(type(squares(11))) == "<class 'list'>", "type of squares(11) is not list"
assert str(len(squares(11))) == "11", "length of squares(11) is not correct"
assert str(sorted([str(element) for element in squares(11)])) == "['1', '100', '121', '16', '25', '36', '4', '49', '64', '81', '9']", "values of squares(11) are not correct"
assert str(squares(11)) == "[1, 4, 9, 16, 25, 36, 49, 64, 81, 100, 121]", "order of elements of squares(11) is not correct"

assert str(type(3)) == "<class 'int'>", "type of 3 is not int"
assert str(3) == "3", "value of 3 is not correct"

assert str(type(squares(3))) == "<class 'list'>", "type of squares(3) is not list"
assert str(len(squares(3))) == "3", "length of squares(3) is not correct"
assert str(sorted([str(element) for element in squares(3)])) == "['1', '4', '9']", "values of squares(3) are not correct"
assert str(squares(3)) == "[1, 4, 9]", "order of elements of squares(3) is not correct"

assert sha1(str(type(squares(3))).encode("utf-8")+b"2cdff3e1a76cfbbf").hexdigest() == "1a922d5a4f918ebf17785f547a58a7a4474208c7", "type of squares(3) is not list"
assert sha1(str(len(squares(3))).encode("utf-8")+b"2cdff3e1a76cfbbf").hexdigest() == "619fca0c36241a703bb1612a0ac3801643b889fe", "length of squares(3) is not correct"
assert sha1(str(sorted([str(element) for element in squares(3)])).encode("utf-8")+b"2cdff3e1a76cfbbf").hexdigest() == "d5e6762d7a5494c9c7626293e5d381cb8572c6ae", "values of squares(3) are not correct"
assert sha1(str(squares(3)).encode("utf-8")+b"2cdff3e1a76cfbbf").hexdigest() == "077e0f4197b1a2372bb39250223a47a7ad8dbf80", "order of elements of squares(3) is not correct"

print('Success!')
```

<!-- grade_id: squares_invalid_input, points: 1, grade: true, solution: false -->
```python
"""Check that squares raises an error for invalid inputs"""
def test_func_throws(func, ErrorType):
    try:
        func()
    except ErrorType:
        return True
    else:
        print('Did not raise right type of error!')
        return False
    
from hashlib import sha1
assert str(type(test_func_throws(lambda : squares(0), ValueError))) == "<class 'bool'>", "type of test_func_throws(lambda : squares(0), ValueError) is not bool"
assert str(test_func_throws(lambda : squares(0), ValueError)) == "True", "boolean value of test_func_throws(lambda : squares(0), ValueError) is not correct"

assert str(type(test_func_throws(lambda : squares(-4), ValueError))) == "<class 'bool'>", "type of test_func_throws(lambda : squares(-4), ValueError) is not bool"
assert str(test_func_throws(lambda : squares(-4), ValueError)) == "True", "boolean value of test_func_throws(lambda : squares(-4), ValueError) is not correct"


print('Success!')
```

---

## Part B (1 point)

Using your `squares` function, write a function that computes the sum of the squares of the numbers from 1 to $n$. Your function should call the `squares` function -- it should NOT reimplement its functionality.

<!-- grade_id: sum_of_squares, points: 0, grade: false, solution: true -->
```python
def sum_of_squares(n):
    """Compute the sum of the squares of numbers from 1 to n."""
    ### BEGIN SOLUTION
    return sum(squares(n))
    ### END SOLUTION
```

The sum of squares from 1 to 10 should be 385. Verify that this is the answer you get:

```python
sum_of_squares(10)
```

<!-- grade_id: correct_sum_of_squares, points: 0.5, grade: true, solution: false -->
```python
"""Check that sum_of_squares returns the correct answer for various inputs."""
from hashlib import sha1
assert str(type(sum_of_squares(1))) == "<class 'int'>", "type of sum_of_squares(1) is not int"
assert str(sum_of_squares(1)) == "1", "value of sum_of_squares(1) is not correct"

assert str(type(sum_of_squares(2))) == "<class 'int'>", "type of sum_of_squares(2) is not int"
assert str(sum_of_squares(2)) == "5", "value of sum_of_squares(2) is not correct"

assert str(type(sum_of_squares(10))) == "<class 'int'>", "type of sum_of_squares(10) is not int"
assert str(sum_of_squares(10)) == "385", "value of sum_of_squares(10) is not correct"

assert str(type(sum_of_squares(11))) == "<class 'int'>", "type of sum_of_squares(11) is not int"
assert str(sum_of_squares(11)) == "506", "value of sum_of_squares(11) is not correct"


print('Success!')
```

<!-- grade_id: sum_of_squares_uses_squares, points: 0.5, grade: true, solution: false -->
```python
"""Check that sum_of_squares relies on squares."""

orig_squares = squares
del squares

from hashlib import sha1
assert str(type(test_func_throws(lambda : sum_of_squares(1), NameError))) == "<class 'bool'>", "type of test_func_throws(lambda : sum_of_squares(1), NameError) is not bool"
assert str(test_func_throws(lambda : sum_of_squares(1), NameError)) == "True", "boolean value of test_func_throws(lambda : sum_of_squares(1), NameError) is not correct"


squares = orig_squares

print('Success!')
```

---
## Part C (1 point)

Using LaTeX math notation, write out the equation that is implemented by your `sum_of_squares` function.

<!-- grade_id: sum_of_squares_equation, points: 1, grade: true, solution: true -->
$\sum_{i=1}^n i^2$

---
## Part D (2 points)

Find a usecase for your `sum_of_squares` function and implement that usecase in the cell below.

<!-- grade_id: sum_of_squares_application, points: 2, grade: true, solution: true -->
```python
def pyramidal_number(n):
    """Returns the n^th pyramidal number"""
    return sum_of_squares(n)
```

---
## Part E (4 points)

State the formulae for an arithmetic and geometric sum and verify them numerically for an example of your choice.