### str

<!-- grade_id: cell-5db49a51603c3906, points: 0, grade: false, solution: true -->
```python
# sample_string = 'Hello'

# YOUR CODE HERE
answers = [1, 'hi', 'bello', 'hello', 'Hello']
sample_string = answers[4]
```

<!-- grade_id: cell-8ce1c8adb2dc6ec3, points: 1, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_string)) == "<class 'str'>", "type is not str"
assert str(len(sample_string)) == "5", "length is not correct"
assert str(sample_string.lower()) == "hello", "value is not correct"
assert str(sample_string) == "Hello", "correct string but incorrect case of letters"

assert sha1(str(type(sample_string)).encode("utf-8")+b"e61eb63245cf74cf").hexdigest() == "f1de95f086601ae85529f2063cd90d81f4f25d4b", "type is not str"
assert sha1(str(len(sample_string)).encode("utf-8")+b"e61eb63245cf74cf").hexdigest() == "1b1470bc92c3e86b2bedf186fe525362e30572bf", "length is not correct"
assert sha1(str(sample_string.lower()).encode("utf-8")+b"e61eb63245cf74cf").hexdigest() == "be6cc87020607b0bd97d463894bcbda71a7e45ad", "value is not correct"
assert sha1(str(sample_string).encode("utf-8")+b"e61eb63245cf74cf").hexdigest() == "c1c7628e7baa42cca639f32c2f5af7e8070a0ba8", "correct string but incorrect case of letters"

print('Success!')
```

Output:
```
Success!
```

### int

<!-- grade_id: cell-20ef3939e635dabb, points: 0, grade: false, solution: true -->
```python
# sample_int = 35

# YOUR CODE HERE
answers = ['hi', 2.4, 2, 35]
sample_int = answers[3]
```

<!-- grade_id: cell-138112d9b5ae4caa, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_int)) == "<class 'int'>", "type is not int"
assert str(sample_int) == "35", "value is not correct."

assert sha1(str(type(sample_int)).encode("utf-8")+b"488dc9e2d44623db").hexdigest() == "da233930b01911d7ac586a8bd321fdcd3de2fbe7", "type is not int"
assert sha1(str(sample_int).encode("utf-8")+b"488dc9e2d44623db").hexdigest() == "f96ea79805a990df951c1507a949f550a6ad4743", "value is not correct."

print('Success!')
```

Output:
```
Success!
```

### float

<!-- grade_id: cell-0ad2fc0d71a06f59, points: 0, grade: false, solution: true -->
```python
# sample_float = 35.98765345

# YOUR CODE HERE
answers = ['hi', 35, 35.98, 35.98765300]
sample_float = answers[0]
```

<!-- grade_id: cell-126da5a3faa851f1, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_float)) == "<class 'int'>", "type is not int"
assert str(sample_float) == "35", "value is not correct."

assert sha1(str(type(sample_float)).encode("utf-8")+b"91ea0919dab99ada").hexdigest() == "7d25ae9bfc41d8c090fcb0a75d3d68940a00c25c", "type is not int"
assert sha1(str(sample_float).encode("utf-8")+b"91ea0919dab99ada").hexdigest() == "91ddb426ed278c2ede362d3bf52d14746c018e22", "value is not correct."

print('Success!')
```

Output:
```
AssertionError                            Traceback (most recent call last)
C:\Users\ALIREZ~1.IRA\AppData\Local\Temp/ipykernel_15564/2254524.py in <cell line: 2>()
      1 from hashlib import sha1
----> 2 assert str(type(sample_float)) == "<class 'int'>", "type is not int"
      3 assert str(sample_float) == "35", "value is not correct."
      4 
      5 assert sha1(str(type(sample_float)).encode("utf-8")+b"91ea0919dab99ada").hexdigest() == "7d25ae9bfc41d8c090fcb0a75d3d68940a00c25c", "type is not int"

AssertionError: type is not int
```

### list

<!-- grade_id: cell-43678e65bae42fd4, points: 0, grade: false, solution: true -->
```python
# sample_list = [1, 'b', "a", 4.5]

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c474a81d248cbe37, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_list)) == "<class 'list'>", "type is not list"
assert str(len(sample_list)) == "4", "length is not correct"
assert str(sorted([str(element) for element in sample_list])) == "['1', '4.5', 'a', 'b']", "values are not correct"
assert str(sample_list) == "[1, 'b', 'a', 4.5]", "order of elements are not correct"

assert sha1(str(type(sample_list)).encode("utf-8")+b"f2050840a5856ff0").hexdigest() == "00df8daa4ca675830d621b74f0702c75ea61a193", "type is not list"
assert sha1(str(len(sample_list)).encode("utf-8")+b"f2050840a5856ff0").hexdigest() == "66562e6e98278c1e3bbf2676ba4b76b582734f8f", "length is not correct"
assert sha1(str(sorted([str(element) for element in sample_list])).encode("utf-8")+b"f2050840a5856ff0").hexdigest() == "d6012e4d0018da058d3e2f4097ef3e187dad8b8c", "values are not correct"
assert sha1(str(sample_list).encode("utf-8")+b"f2050840a5856ff0").hexdigest() == "691052a0ce3e5900dcf3ec927d2786ed7655b2fc", "order of elements are not correct"

print('Success!')
```

### set

<!-- grade_id: cell-fe5861008f297022, points: 0, grade: false, solution: true -->
```python
# sample_set = {1, 2, 3, 4, 5}

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-5e1c667195b47187, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_set)) == "<class 'set'>", "type is not set"
assert str(len(sample_set)) == "5", "length is not correct"
assert str(sample_set) == "{1, 2, 3, 4, 5}", "value is not correct"

assert sha1(str(type(sample_set)).encode("utf-8")+b"3397cc2d891b345a").hexdigest() == "e186f70edcbda744f415970baffe634a6b9671dd", "type is not set"
assert sha1(str(len(sample_set)).encode("utf-8")+b"3397cc2d891b345a").hexdigest() == "38ed33fa27697456cdf65113df851e3f256e5748", "length is not correct"
assert sha1(str(sample_set).encode("utf-8")+b"3397cc2d891b345a").hexdigest() == "fee0a8690d3862595fa07899c1b03852cc3dbdb0", "value is not correct"

print('Success!')
```

### tuple

<!-- grade_id: cell-a971551f5f91de4b, points: 0, grade: false, solution: true -->
```python
# sample_tuple = (23, 78)

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-48cfa5ac37cdd98b, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_tuple)) == "<class 'tuple'>", "type is not tuple"
assert str(len(sample_tuple)) == "2", "length is not correct"
assert str(sorted([str(element) for element in sample_tuple])) == "['23', '78']", "values are not correct"
assert str(sample_tuple) == "(23, 78)", "order of elements are not correct"

assert sha1(str(type(sample_tuple)).encode("utf-8")+b"0a720c80aa4cb76f").hexdigest() == "27f23d5eadcbd912f84dcef9a25cdf3a7d277a7c", "type is not tuple"
assert sha1(str(len(sample_tuple)).encode("utf-8")+b"0a720c80aa4cb76f").hexdigest() == "74c47a81a4b7608f719c48c8509c3627edf6c622", "length is not correct"
assert sha1(str(sorted([str(element) for element in sample_tuple])).encode("utf-8")+b"0a720c80aa4cb76f").hexdigest() == "a91025e0ebc0fc4aac150b26478e8e844ae2fb71", "values are not correct"
assert sha1(str(sample_tuple).encode("utf-8")+b"0a720c80aa4cb76f").hexdigest() == "aa62b7d22c3eda8a2159b51d2bb93dde762c10c7", "order of elements are not correct"

print('Success!')
```

### dict

<!-- grade_id: cell-ddabbd9234db36c0, points: 0, grade: false, solution: true -->
```python
# sample_dict = {'a': 1, 'b': 2, "c": 3}

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b8b1ad86fbbe578f, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_dict)) == "<class 'dict'>", "type is not dict"
assert str(len(list(sample_dict.keys()))) == "3", "number of keys is not correct"
assert str(sorted(list(sample_dict.keys()))) == "['a', 'b', 'c']", "keys are not correct"
assert str(sorted(list(sample_dict.values()))) == "[1, 2, 3]", "values are not correct"
assert str(sample_dict) == "{'a': 1, 'b': 2, 'c': 3}", "not correct"

assert sha1(str(type(sample_dict)).encode("utf-8")+b"37c86b6aed9bbf3c").hexdigest() == "d8a1fb1de504b234463de5318291e4dc7d4be289", "type is not dict"
assert sha1(str(len(list(sample_dict.keys()))).encode("utf-8")+b"37c86b6aed9bbf3c").hexdigest() == "9d700bb36272a8d8f29067b2fff3e2a85a21b7f3", "number of keys is not correct"
assert sha1(str(sorted(list(sample_dict.keys()))).encode("utf-8")+b"37c86b6aed9bbf3c").hexdigest() == "6f71976d508e97ac4347cd76a967fcdc1ab59bd0", "keys are not correct"
assert sha1(str(sorted(list(sample_dict.values()))).encode("utf-8")+b"37c86b6aed9bbf3c").hexdigest() == "f9a0804430c9597e531da0d8cfa876fad194cba8", "values are not correct"
assert sha1(str(sample_dict).encode("utf-8")+b"37c86b6aed9bbf3c").hexdigest() == "b1cbd93dd454c8743fdac3f4c0440b932bcc5387", "not correct"

print('Success!')
```

### bool

<!-- grade_id: cell-e807929d33f8deb6, points: 0, grade: false, solution: true -->
```python
# sample_bool = True

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-cf2f289017808ac8, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_bool)) == "<class 'bool'>", "type is not bool"
assert str(sample_bool) == "True", "boolean value is not correct"

assert sha1(str(type(sample_bool)).encode("utf-8")+b"3a285f02396eda30").hexdigest() == "f1a13b5ca2b9ea6a2a29c8990e71fa09df1c03d5", "type is not bool"
assert sha1(str(sample_bool).encode("utf-8")+b"3a285f02396eda30").hexdigest() == "b005a9aa6b9707e6c0d58a9e8dc5814ee15613f7", "boolean value is not correct"

print('Success!')
```

```python

```