### str

<!-- grade_id: cell-5db49a51603c3906, points: 0, grade: false, solution: true -->
```python
# sample_string = 'Hello'

# YOUR CODE HERE
answers = [1, 'hi', 'bello', 'hello', 'Hello']
sample_string = answers[4]
```

<!-- grade_id: cell-8ce1c8adb2dc6ec3, points: 1, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_string)) == "<class 'str'>", "type is not str"
assert str(len(sample_string)) == "5", "length is not correct"
assert str(sample_string.lower()) == "hello", "value is not correct"
assert str(sample_string) == "Hello", "correct string but incorrect case of letters"

assert sha1(str(type(sample_string)).encode("utf-8")+b"1996c775d48999dd").hexdigest() == "cde20e4d019e2db1a90178690a215263a4fe10f9", "type is not str"
assert sha1(str(len(sample_string)).encode("utf-8")+b"1996c775d48999dd").hexdigest() == "2356178c506d45dc7667c1081c68c1127d5ea59c", "length is not correct"
assert sha1(str(sample_string.lower()).encode("utf-8")+b"1996c775d48999dd").hexdigest() == "644344c88b1cadda335c62aedb604597dfbbcabe", "value is not correct"
assert sha1(str(sample_string).encode("utf-8")+b"1996c775d48999dd").hexdigest() == "08ce6beb13f255ca0a369e757eafc38e41910a0e", "correct string but incorrect case of letters"

print('Success!')
```

### int

<!-- grade_id: cell-20ef3939e635dabb, points: 0, grade: false, solution: true -->
```python
# sample_int = 35

# YOUR CODE HERE
answers = ['hi', 2.4, 2, 35]
sample_int = answers[3]
```

<!-- grade_id: cell-138112d9b5ae4caa, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_int)) == "<class 'int'>", "type is not int"
assert str(sample_int) == "35", "value is not correct"

assert sha1(str(type(sample_int)).encode("utf-8")+b"647116cd3145be4a").hexdigest() == "ccf5c2cc34f88080427d6fbc63f79200e8a08bbe", "type is not int"
assert sha1(str(sample_int).encode("utf-8")+b"647116cd3145be4a").hexdigest() == "bc6952c2db4cfb5b8df7610aafb933c1befc3b55", "value is not correct"

print('Success!')
```

### float

<!-- grade_id: cell-0ad2fc0d71a06f59, points: 0, grade: false, solution: true -->
```python
# sample_float = 35.98765345

# YOUR CODE HERE
answers = ['hi', 35, 35.98, 35.98765300]
sample_float = answers[3]
```

<!-- grade_id: cell-126da5a3faa851f1, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_float)) == "<class 'float'>", "type is not float"
assert str(round(sample_float, 5)) == "35.98765", "value neither correct nor close"

assert sha1(str(type(sample_float)).encode("utf-8")+b"073d7ae4dd1d8397").hexdigest() == "651aa63102af931443aac11d18ac929c86556146", "type is not float"
assert sha1(str(round(sample_float, 5)).encode("utf-8")+b"073d7ae4dd1d8397").hexdigest() == "8bfd1b4ed2a7ef3cf7ad4f4e886541b05cd2d8da", "value neither correct nor close"

print('Success!')
```

Output:
```
Success!
```

### list

<!-- grade_id: cell-43678e65bae42fd4, points: 0, grade: false, solution: true -->
```python
# sample_list = [1, 'b', "a", 4.5]

# YOUR CODE HERE
answers = [[1, 'B', "a", 4.5], [1, "a", 'b', 4.5], [1, 'b', "a", 4.5]]
sample_list = answers[2]
```

<!-- grade_id: cell-c474a81d248cbe37, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_list)) == "<class 'list'>", "type is not list"
assert str(len(sample_list)) == "4", "length is not correct"
assert str(sorted([str(element) for element in sample_list])) == "['1', '4.5', 'a', 'b']", "values are not correct"
assert str(sample_list) == "[1, 'b', 'a', 4.5]", "order of elements are not correct"

assert sha1(str(type(sample_list)).encode("utf-8")+b"e8d852479eff7351").hexdigest() == "e55b4f21e3150a793780124d14222d626fb33b32", "type is not list"
assert sha1(str(len(sample_list)).encode("utf-8")+b"e8d852479eff7351").hexdigest() == "ba3b3b98eff96d1458c420929645591b8f7b3dd1", "length is not correct"
assert sha1(str(sorted([str(element) for element in sample_list])).encode("utf-8")+b"e8d852479eff7351").hexdigest() == "36ed30d18027613e0ae7885acedd3e852e8756f0", "values are not correct"
assert sha1(str(sample_list).encode("utf-8")+b"e8d852479eff7351").hexdigest() == "3ae6fd42834d82cf292fdaf3e37f71f0d75837da", "order of elements are not correct"

print('Success!')
```

Output:
```
Success!
```

### set

<!-- grade_id: cell-fe5861008f297022, points: 0, grade: false, solution: true -->
```python
# sample_set = {1, 2, 3, 4, 5}

# YOUR CODE HERE
answers = [{1, 2, 3, 4, 5, 5}, {1, 2, 3, 5, 4}, {1, 'b', "a", 4.5}, {1, 2, 3, 5, 'a'}]
sample_set = answers[0]
```

<!-- grade_id: cell-5e1c667195b47187, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_set)) == "<class 'set'>", "type is not set"
assert str(len(sample_set)) == "5", "length is not correct"
assert str(sample_set) == "{1, 2, 3, 4, 5}", "value is not correct"

assert sha1(str(type(sample_set)).encode("utf-8")+b"55726ed71fc4ba04").hexdigest() == "687634aa3e21cd359b3e1aea3133844a026f480e", "type is not set"
assert sha1(str(len(sample_set)).encode("utf-8")+b"55726ed71fc4ba04").hexdigest() == "3f10873af8a64dd7e366a8c5a96eac85c69a9b02", "length is not correct"
assert sha1(str(sample_set).encode("utf-8")+b"55726ed71fc4ba04").hexdigest() == "25a2c594d3a9eaecdd84fff3fccb8d61b2714768", "value is not correct"

print('Success!')
```

Output:
```
Success!
```

### tuple

<!-- grade_id: cell-a971551f5f91de4b, points: 0, grade: false, solution: true -->
```python
# sample_tuple = (23, 78)

# YOUR CODE HERE
answers = ['a', (23, 78, 34), (78, 23), (23, 'a'), (23, 78)]
sample_tuple = answers[4]
```

<!-- grade_id: cell-48cfa5ac37cdd98b, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_tuple)) == "<class 'tuple'>", "type is not tuple"
assert str(len(sample_tuple)) == "2", "length is not correct"
assert str(sorted([str(element) for element in sample_tuple])) == "['23', '78']", "values are not correct"
assert str(sample_tuple) == "(23, 78)", "order of elements are not correct"

assert sha1(str(type(sample_tuple)).encode("utf-8")+b"41d40d614e5158e1").hexdigest() == "97afad47bc6a9b26977f3fc0bff4ca01b72bb9d5", "type is not tuple"
assert sha1(str(len(sample_tuple)).encode("utf-8")+b"41d40d614e5158e1").hexdigest() == "d416498af3315d7ded644e9a4ec7bba4a9617678", "length is not correct"
assert sha1(str(sorted([str(element) for element in sample_tuple])).encode("utf-8")+b"41d40d614e5158e1").hexdigest() == "943391189bf8c06acaa1f618cc4e5d4fbedd740c", "values are not correct"
assert sha1(str(sample_tuple).encode("utf-8")+b"41d40d614e5158e1").hexdigest() == "518cf96875c6d3de3a76cf11f68c155178acf3a7", "order of elements are not correct"

print('Success!')
```

Output:
```
Success!
```

### dict

<!-- grade_id: cell-ddabbd9234db36c0, points: 0, grade: false, solution: true -->
```python
# sample_dict = {'a': 1, 'b': 2, "c": 3}

# YOUR CODE HERE
answers = ['a', {}, {'a': 1, 'b': 2, "c": 3, 'd': 4}, {'a': 1, 'c': 2, "b": 3}]
sample_dict = answers[3]
```

<!-- grade_id: cell-b8b1ad86fbbe578f, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_dict)) == "<class 'dict'>", "type is not dict"
assert str(len(list(sample_dict.keys()))) == "3", "number of keys is not correct"
assert str(sorted(list(sample_dict.keys()))) == "['a', 'b', 'c']", "keys are not correct"
assert str(sorted(list(sample_dict.values()))) == "[1, 2, 3]", "values are not correct"
assert str(sample_dict) == "{'a': 1, 'b': 2, 'c': 3}", "not correct"

assert sha1(str(type(sample_dict)).encode("utf-8")+b"c337926e481ea9b2").hexdigest() == "46e681dea3b2a6346c354d1dcea78603bac3d823", "type is not dict"
assert sha1(str(len(list(sample_dict.keys()))).encode("utf-8")+b"c337926e481ea9b2").hexdigest() == "0a110352b96589eb902896beaa46dff36e07e267", "number of keys is not correct"
assert sha1(str(sorted(list(sample_dict.keys()))).encode("utf-8")+b"c337926e481ea9b2").hexdigest() == "d11e08f9f617348d36f9d7db561bc6a78507fbb2", "keys are not correct"
assert sha1(str(sorted(list(sample_dict.values()))).encode("utf-8")+b"c337926e481ea9b2").hexdigest() == "e1949fdfce30fbb185817893a12f8c1a1ddba3ff", "values are not correct"
assert sha1(str(sample_dict).encode("utf-8")+b"c337926e481ea9b2").hexdigest() == "984574a12dc64801aab3fb4dfb7ad082ceebbe56", "not correct"

print('Success!')
```

Output:
```
AssertionError                            Traceback (most recent call last)
C:\Users\ALIREZ~1.IRA\AppData\Local\Temp/ipykernel_15732/2191625916.py in <cell line: 6>()
      4 assert str(sorted(list(sample_dict.keys()))) == "['a', 'b', 'c']", "keys are not correct"
      5 assert str(sorted(list(sample_dict.values()))) == "[1, 2, 3]", "values are not correct"
----> 6 assert str(sample_dict) == "{'a': 1, 'b': 2, 'c': 3}", "not correct"
      7 
      8 assert sha1(str(type(sample_dict)).encode("utf-8")+b"c337926e481ea9b2").hexdigest() == "46e681dea3b2a6346c354d1dcea78603bac3d823", "type is not dict"

AssertionError: not correct
```

### bool

<!-- grade_id: cell-e807929d33f8deb6, points: 0, grade: false, solution: true -->
```python
# sample_bool = True

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-cf2f289017808ac8, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_bool)) == "<class 'bool'>", "type is not bool"
assert str(sample_bool) == "True", "boolean value is not correct"

assert sha1(str(type(sample_bool)).encode("utf-8")+b"6fc511d5e4c1b9d4").hexdigest() == "aab11d51049f81fb8d062f4d829dd68ff5635828", "type is not bool"
assert sha1(str(sample_bool).encode("utf-8")+b"6fc511d5e4c1b9d4").hexdigest() == "e7ca76a78d6906acc6aeaeae79becbcd2aedee5c", "boolean value is not correct"

print('Success!')
```

```python

```