### str

<!-- grade_id: cell-5db49a51603c3906, points: 0, grade: false, solution: true -->
```python
# sample_string = 'Hello'

### BEGIN SOLUTION
sample_string = 'Hello'
### END SOLUTION
```

<!-- grade_id: cell-8ce1c8adb2dc6ec3, points: 1, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_string)) == "<class 'str'>", "type of sample_string is not str"
assert str(len(sample_string)) == "5", "length of sample_string is not correct"
assert str(sample_string.lower()) == "hello", "value of sample_string is not correct"
assert str(sample_string) == "Hello", "correct string value of sample_string but incorrect case of letters"

assert sha1(str(type(sample_string)).encode("utf-8")+b"861eac6bd29a604c").hexdigest() == "52df73e14f67de23f89d4ce0074aaec8d7d90e68", "type of sample_string is not str"
assert sha1(str(len(sample_string)).encode("utf-8")+b"861eac6bd29a604c").hexdigest() == "c8c5340685d1ede63cbc2cdd6bd4a1ac5636d12b", "length of sample_string is not correct"
assert sha1(str(sample_string.lower()).encode("utf-8")+b"861eac6bd29a604c").hexdigest() == "e470700bf6d3248e4ab3489f377cac1f150d4cdf", "value of sample_string is not correct"
assert sha1(str(sample_string).encode("utf-8")+b"861eac6bd29a604c").hexdigest() == "c544e601bcfb7271954f9eab8d21e1bc692bf1bc", "correct string value of sample_string but incorrect case of letters"

print('Success!')
```

### int

<!-- grade_id: cell-20ef3939e635dabb, points: 0, grade: false, solution: true -->
```python
# sample_int = 35

### BEGIN SOLUTION
sample_int = 35
### END SOLUTION
```

<!-- grade_id: cell-138112d9b5ae4caa, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_int)) == "<class 'int'>", "type of sample_int is not int"
assert str(sample_int) == "35", "value of sample_int is not correct"

assert sha1(str(type(sample_int)).encode("utf-8")+b"255e03f247247a6e").hexdigest() == "f730d388de4a87a9421cd304cdfdda6a1857f2ab", "type of sample_int is not int"
assert sha1(str(sample_int).encode("utf-8")+b"255e03f247247a6e").hexdigest() == "6ef0f73b2e226a191bdab80be8ee0837acbd141e", "value of sample_int is not correct"

print('Success!')
```

### float

<!-- grade_id: cell-0ad2fc0d71a06f59, points: 0, grade: false, solution: true -->
```python
# sample_float = 35.98765345

### BEGIN SOLUTION
sample_float = 35.98765345
### END SOLUTION
```

<!-- grade_id: cell-126da5a3faa851f1, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_float)) == "<class 'float'>", "type of sample_float is not float"
assert str(round(sample_float, 8)) == "35.98765345", "value of sample_float is not correct (rounded to 8 decimal places)"

assert sha1(str(type(sample_float)).encode("utf-8")+b"6c1645c48088af85").hexdigest() == "1b445f37cb3e039e71a3d445ed2da568e3a28dbb", "type of sample_float is not float"
assert sha1(str(round(sample_float, 8)).encode("utf-8")+b"6c1645c48088af85").hexdigest() == "1ab65b6dbea5ef62e9936818e807d9b924579548", "value of sample_float is not correct (rounded to 8 decimal places)"

print('Success!')
```

### list

<!-- grade_id: cell-43678e65bae42fd4, points: 0, grade: false, solution: true -->
```python
# sample_list = [1, 'b', "a", 4.5]

### BEGIN SOLUTION
sample_list = [1, 'b', "a", 4.5]
### END SOLUTION
```

<!-- grade_id: cell-c474a81d248cbe37, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_list)) == "<class 'list'>", "type of sample_list is not list"
assert str(len(sample_list)) == "4", "length of sample_list is not correct"
assert str(sorted([str(element) for element in sample_list])) == "['1', '4.5', 'a', 'b']", "values of sample_list are not correct"
assert str(sample_list) == "[1, 'b', 'a', 4.5]", "order of elements of sample_list is not correct"

assert sha1(str(type(sample_list)).encode("utf-8")+b"c9e79fb58deb6fec").hexdigest() == "d7c873ecfa2086f173b3827a7b88465b15d609fd", "type of sample_list is not list"
assert sha1(str(len(sample_list)).encode("utf-8")+b"c9e79fb58deb6fec").hexdigest() == "8d7e89aa5fd40606921f36da42da85cdade4b619", "length of sample_list is not correct"
assert sha1(str(sorted([str(element) for element in sample_list])).encode("utf-8")+b"c9e79fb58deb6fec").hexdigest() == "4b1c5fe5ad82a6c162cb671d144bedb4cfb2e55a", "values of sample_list are not correct"
assert sha1(str(sample_list).encode("utf-8")+b"c9e79fb58deb6fec").hexdigest() == "0617516f46aeaad1f4840c92c309a77711a23acf", "order of elements of sample_list is not correct"

print('Success!')
```

### set

<!-- grade_id: cell-fe5861008f297022, points: 0, grade: false, solution: true -->
```python
# sample_set = {1, 2, 3, 4, 5}

### BEGIN SOLUTION
sample_set = {1, 2, 3, 4, 5}
### END SOLUTION
```

<!-- grade_id: cell-5e1c667195b47187, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_set)) == "<class 'set'>", "type of sample_set is not set"
assert str(len(sample_set)) == "5", "length of sample_set is not correct"
assert str(sample_set) == "{1, 2, 3, 4, 5}", "value of sample_set is not correct"

assert sha1(str(type(sample_set)).encode("utf-8")+b"edde70af755b84b3").hexdigest() == "04dc83a79d71f7fa82273d428821a3f04d90210d", "type of sample_set is not set"
assert sha1(str(len(sample_set)).encode("utf-8")+b"edde70af755b84b3").hexdigest() == "6b506d52da55a71de1be0941205d501738618487", "length of sample_set is not correct"
assert sha1(str(sample_set).encode("utf-8")+b"edde70af755b84b3").hexdigest() == "91915a1ad1fe52804a25648b12e4964a232f5e1a", "value of sample_set is not correct"

print('Success!')
```

### tuple

<!-- grade_id: cell-a971551f5f91de4b, points: 0, grade: false, solution: true -->
```python
# sample_tuple = (23, 78)

### BEGIN SOLUTION
sample_tuple = (23, 78)
### END SOLUTION
```

<!-- grade_id: cell-48cfa5ac37cdd98b, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_tuple)) == "<class 'tuple'>", "type of sample_tuple is not tuple"
assert str(len(sample_tuple)) == "2", "length of sample_tuple is not correct"
assert str(sorted([str(element) for element in sample_tuple])) == "['23', '78']", "values of sample_tuple are not correct"
assert str(sample_tuple) == "(23, 78)", "order of elements of sample_tuple is not correct"

assert sha1(str(type(sample_tuple)).encode("utf-8")+b"5aa0dade2c0537a4").hexdigest() == "a2824e1d217368af581a99bf527e9a0f7f8862ad", "type of sample_tuple is not tuple"
assert sha1(str(len(sample_tuple)).encode("utf-8")+b"5aa0dade2c0537a4").hexdigest() == "70b1ad2b4da3848d0bc2dd23f0c276958739b144", "length of sample_tuple is not correct"
assert sha1(str(sorted([str(element) for element in sample_tuple])).encode("utf-8")+b"5aa0dade2c0537a4").hexdigest() == "5c2c9cf324c9f6390e4438ead8d747e7cdfd57ca", "values of sample_tuple are not correct"
assert sha1(str(sample_tuple).encode("utf-8")+b"5aa0dade2c0537a4").hexdigest() == "ab75f8b582b497d0647a5079107dc667bf56ebe9", "order of elements of sample_tuple is not correct"

print('Success!')
```

### dict

<!-- grade_id: cell-ddabbd9234db36c0, points: 0, grade: false, solution: true -->
```python
# sample_dict = {'a': 1, 'b': 2, "c": 3}

### BEGIN SOLUTION
sample_dict = {'a': 1, 'b': 2, "c": 3}
### END SOLUTION
```

<!-- grade_id: cell-b8b1ad86fbbe578f, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_dict)) == "<class 'dict'>", "type of sample_dict is not dict"
assert str(len(list(sample_dict.keys()))) == "3", "number of keys of sample_dict is not correct"
assert str(sorted(list(sample_dict.keys()))) == "['a', 'b', 'c']", "keys of sample_dict are not correct"
assert str(sorted(list(sample_dict.values()))) == "[1, 2, 3]", "correct keys, but values of sample_dict are not correct"
assert str(sample_dict) == "{'a': 1, 'b': 2, 'c': 3}", "correct keys and values, but incorrect correspondence in keys and values of sample_dict"

assert sha1(str(type(sample_dict)).encode("utf-8")+b"f8c4918c597e5958").hexdigest() == "4641d2c586abf6f03531f474a264def1973c453a", "type of sample_dict is not dict"
assert sha1(str(len(list(sample_dict.keys()))).encode("utf-8")+b"f8c4918c597e5958").hexdigest() == "3b12f1be84319d234bc2cc3638e473390811070a", "number of keys of sample_dict is not correct"
assert sha1(str(sorted(list(sample_dict.keys()))).encode("utf-8")+b"f8c4918c597e5958").hexdigest() == "7b7f2228bb060840f171f34e200dfe70f9d62750", "keys of sample_dict are not correct"
assert sha1(str(sorted(list(sample_dict.values()))).encode("utf-8")+b"f8c4918c597e5958").hexdigest() == "06773318d16784799daaaa2d7b75515531292b9c", "correct keys, but values of sample_dict are not correct"
assert sha1(str(sample_dict).encode("utf-8")+b"f8c4918c597e5958").hexdigest() == "2c559920542a9eae1afdf1de1baa43cdffd16cb0", "correct keys and values, but incorrect correspondence in keys and values of sample_dict"

print('Success!')
```

### bool

<!-- grade_id: cell-e807929d33f8deb6, points: 0, grade: false, solution: true -->
```python
# sample_bool = True

### BEGIN SOLUTION
sample_bool = True
### END SOLUTION
```

<!-- grade_id: cell-cf2f289017808ac8, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_bool)) == "<class 'bool'>", "type of sample_bool is not bool"
assert str(sample_bool) == "True", "boolean value of sample_bool is not correct"

assert sha1(str(type(sample_bool)).encode("utf-8")+b"cd5d1b21c86535cf").hexdigest() == "b2392b92f7abda9bfaa56353b62f2ecba9621787", "type of sample_bool is not bool"
assert sha1(str(sample_bool).encode("utf-8")+b"cd5d1b21c86535cf").hexdigest() == "4e736129cdad6ee0bace988e9ed2b82b4ebe5465", "boolean value of sample_bool is not correct"

print('Success!')
```

```python

```