### str

<!-- grade_id: cell-5db49a51603c3906, points: 0, grade: false, solution: true -->
```python
# sample_string = 'Hello'

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8ce1c8adb2dc6ec3, points: 1, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_string)) == "<class 'str'>", "type is not str"
assert str(len(sample_string)) == "5", "length is not correct"
assert str(sample_string.lower()) == "hello", "value is not correct"
assert str(sample_string) == "Hello", "correct string but incorrect case of letters"

assert sha1(str(type(sample_string)).encode("utf-8")+b"65eaa70425eede06").hexdigest() == "97c8e1d58a1f93d1b8aa6c75d7a921d404999a02", "type is not str"
assert sha1(str(len(sample_string)).encode("utf-8")+b"65eaa70425eede06").hexdigest() == "85ec56831ba94d766e1717cd7a9f5d3428b0edc4", "length is not correct"
assert sha1(str(sample_string.lower()).encode("utf-8")+b"65eaa70425eede06").hexdigest() == "5756de6a785b8772d44f1e8b0f4f6601ef800c85", "value is not correct"
assert sha1(str(sample_string).encode("utf-8")+b"65eaa70425eede06").hexdigest() == "e18c134d1f3ca405d937adc9c1cbbd98c2dc67aa", "correct string but incorrect case of letters"

print('Success!')
```

### int

<!-- grade_id: cell-20ef3939e635dabb, points: 0, grade: false, solution: true -->
```python
# sample_int = 35

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-138112d9b5ae4caa, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_int)) == "<class 'int'>", "type is not int"
assert str(sample_int) == "35", "value is not correct"

assert sha1(str(type(sample_int)).encode("utf-8")+b"b008f818fd6a5e3d").hexdigest() == "2711b0eb32adcde9afafb140b43999c4fde8568d", "type is not int"
assert sha1(str(sample_int).encode("utf-8")+b"b008f818fd6a5e3d").hexdigest() == "f69cbe2548da941f23b99a2a6f595d47197a61c0", "value is not correct"

print('Success!')
```

### float

<!-- grade_id: cell-0ad2fc0d71a06f59, points: 0, grade: false, solution: true -->
```python
# sample_float = 35.98765345

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-126da5a3faa851f1, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_float)) == "<class 'float'>", "type is not float"
assert str(round(sample_float, 8)) == "35.98765345", "value is not correct (rounded to 8 decimal places)"

assert sha1(str(type(sample_float)).encode("utf-8")+b"295db598760a1cc0").hexdigest() == "f1bfd29f1d6092bc591b8971ff3e79ad9f9a56ed", "type is not float"
assert sha1(str(round(sample_float, 8)).encode("utf-8")+b"295db598760a1cc0").hexdigest() == "36585fbe9fb791d6b90b0eb80a943d7b91220182", "value is not correct (rounded to 8 decimal places)"

print('Success!')
```

### list

<!-- grade_id: cell-43678e65bae42fd4, points: 0, grade: false, solution: true -->
```python
# sample_list = [1, 'b', "a", 4.5]

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c474a81d248cbe37, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_list)) == "<class 'list'>", "type is not list"
assert str(len(sample_list)) == "4", "length is not correct"
assert str(sorted([str(element) for element in sample_list])) == "['1', '4.5', 'a', 'b']", "values are not correct"
assert str(sample_list) == "[1, 'b', 'a', 4.5]", "order of elements is not correct"

assert sha1(str(type(sample_list)).encode("utf-8")+b"72f52e627ab1c089").hexdigest() == "37fe11ec238e62e543cd273ecfce918e009e1056", "type is not list"
assert sha1(str(len(sample_list)).encode("utf-8")+b"72f52e627ab1c089").hexdigest() == "45d7e8513cc82e09c6dd07f8c370d0f9332945a6", "length is not correct"
assert sha1(str(sorted([str(element) for element in sample_list])).encode("utf-8")+b"72f52e627ab1c089").hexdigest() == "3adc4c950a3feb65322e0761146b47980b4545b2", "values are not correct"
assert sha1(str(sample_list).encode("utf-8")+b"72f52e627ab1c089").hexdigest() == "d4ff8d4f8eda395c51647c500525ce79746ec900", "order of elements is not correct"

print('Success!')
```

### set

<!-- grade_id: cell-fe5861008f297022, points: 0, grade: false, solution: true -->
```python
# sample_set = {1, 2, 3, 4, 5}

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-5e1c667195b47187, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_set)) == "<class 'set'>", "type is not set"
assert str(len(sample_set)) == "5", "length is not correct"
assert str(sample_set) == "{1, 2, 3, 4, 5}", "value is not correct"

assert sha1(str(type(sample_set)).encode("utf-8")+b"8f4fb417cd37dbba").hexdigest() == "1637a7c3a0410a666c132a85ad49d7cb45951465", "type is not set"
assert sha1(str(len(sample_set)).encode("utf-8")+b"8f4fb417cd37dbba").hexdigest() == "864cc4be785e55ea9d5e5ec0399af65913db15fe", "length is not correct"
assert sha1(str(sample_set).encode("utf-8")+b"8f4fb417cd37dbba").hexdigest() == "4bbb292abaf4a09ba6d9b1e40390d8597ccae355", "value is not correct"

print('Success!')
```

### tuple

<!-- grade_id: cell-a971551f5f91de4b, points: 0, grade: false, solution: true -->
```python
# sample_tuple = (23, 78)

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-48cfa5ac37cdd98b, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_tuple)) == "<class 'tuple'>", "type is not tuple"
assert str(len(sample_tuple)) == "2", "length is not correct"
assert str(sorted([str(element) for element in sample_tuple])) == "['23', '78']", "values are not correct"
assert str(sample_tuple) == "(23, 78)", "order of elements is not correct"

assert sha1(str(type(sample_tuple)).encode("utf-8")+b"89cdbc8727572eab").hexdigest() == "74aec4872a08dbeb090c79e6c07f61214b47eff9", "type is not tuple"
assert sha1(str(len(sample_tuple)).encode("utf-8")+b"89cdbc8727572eab").hexdigest() == "885a7c2f07091c7366b30c410a355ae02213d5c4", "length is not correct"
assert sha1(str(sorted([str(element) for element in sample_tuple])).encode("utf-8")+b"89cdbc8727572eab").hexdigest() == "6e58ec263c119fa882dcec76da78af5be0b46487", "values are not correct"
assert sha1(str(sample_tuple).encode("utf-8")+b"89cdbc8727572eab").hexdigest() == "4ec68488f49f01018e839a3a109bf9425e2bfc51", "order of elements is not correct"

print('Success!')
```

### dict

<!-- grade_id: cell-ddabbd9234db36c0, points: 0, grade: false, solution: true -->
```python
# sample_dict = {'a': 1, 'b': 2, "c": 3}

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b8b1ad86fbbe578f, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_dict)) == "<class 'dict'>", "type is not dict"
assert str(len(list(sample_dict.keys()))) == "3", "number of keys is not correct"
assert str(sorted(list(sample_dict.keys()))) == "['a', 'b', 'c']", "keys are not correct"
assert str(sorted(list(sample_dict.values()))) == "[1, 2, 3]", "correct keys, but values are not correct"
assert str(sample_dict) == "{'a': 1, 'b': 2, 'c': 3}", "correct keys and values, but incorrect correspondence"

assert sha1(str(type(sample_dict)).encode("utf-8")+b"fa2bc67aceaa9894").hexdigest() == "8e74a042daf13d833042f318d907bb5e9612d384", "type is not dict"
assert sha1(str(len(list(sample_dict.keys()))).encode("utf-8")+b"fa2bc67aceaa9894").hexdigest() == "91b10d410bb81759b8cd31318ba5c8f121a406ca", "number of keys is not correct"
assert sha1(str(sorted(list(sample_dict.keys()))).encode("utf-8")+b"fa2bc67aceaa9894").hexdigest() == "5b52791e238c01c9fea8a11e7e13fbfef86a1351", "keys are not correct"
assert sha1(str(sorted(list(sample_dict.values()))).encode("utf-8")+b"fa2bc67aceaa9894").hexdigest() == "b04f04d8396c25a1b68040892d453144a50d6408", "correct keys, but values are not correct"
assert sha1(str(sample_dict).encode("utf-8")+b"fa2bc67aceaa9894").hexdigest() == "bc7f73f1fa95ad3d4da149354026766e763cfa27", "correct keys and values, but incorrect correspondence"

print('Success!')
```

### bool

<!-- grade_id: cell-e807929d33f8deb6, points: 0, grade: false, solution: true -->
```python
# sample_bool = True

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-cf2f289017808ac8, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_bool)) == "<class 'bool'>", "type is not bool"
assert str(sample_bool) == "True", "boolean value is not correct"

assert sha1(str(type(sample_bool)).encode("utf-8")+b"3626045419a102d5").hexdigest() == "d88b3697328fa3a2a9ae36093b6d3f587bea77c3", "type is not bool"
assert sha1(str(sample_bool).encode("utf-8")+b"3626045419a102d5").hexdigest() == "d4483aeb8a8243e03caef6869a1998175b1df8de", "boolean value is not correct"

print('Success!')
```

```python

```