### str

<!-- grade_id: cell-5db49a51603c3906, points: 0, grade: false, solution: true -->
```python
# sample_string = 'Hello'

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-8ce1c8adb2dc6ec3, points: 1, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_string)) == "<class 'str'>", "type is not str"
assert str(len(sample_string)) == "5", "length is not correct"
assert str(sample_string.lower()) == "hello", "value is not correct"
assert str(sample_string) == "Hello", "correct string but incorrect case of letters"

assert sha1(str(type(sample_string)).encode("utf-8")+b"687a12e0445a85d2").hexdigest() == "270a860037637a3711a65d3c5ab506926fd6fcaa", "type is not str"
assert sha1(str(len(sample_string)).encode("utf-8")+b"687a12e0445a85d2").hexdigest() == "f9dc76019b6d9cba63538871f9492f802ee6d620", "length is not correct"
assert sha1(str(sample_string.lower()).encode("utf-8")+b"687a12e0445a85d2").hexdigest() == "61408b3a686dfc265498cfe39a7b81444f4891cb", "value is not correct"
assert sha1(str(sample_string).encode("utf-8")+b"687a12e0445a85d2").hexdigest() == "2ed136f522b89ff814d7d83b411c58d0fa9f0516", "correct string but incorrect case of letters"

print('Success!')
```

### int

<!-- grade_id: cell-20ef3939e635dabb, points: 0, grade: false, solution: true -->
```python
# sample_int = 35

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-138112d9b5ae4caa, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_int)) == "<class 'int'>", "type is not int"
assert str(sample_int) == "35", "value is not correct"

assert sha1(str(type(sample_int)).encode("utf-8")+b"392526080b2ecf6f").hexdigest() == "f8d4114f158fd0ca0c48078c8dfe89c4858fe869", "type is not int"
assert sha1(str(sample_int).encode("utf-8")+b"392526080b2ecf6f").hexdigest() == "3a4ae7d1216432b7aa39747da23dc42542b173d7", "value is not correct"

print('Success!')
```

### float

<!-- grade_id: cell-0ad2fc0d71a06f59, points: 0, grade: false, solution: true -->
```python
# sample_float = 35.98765345

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-126da5a3faa851f1, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_float)) == "<class 'float'>", "type is not float"
assert str(round(sample_float, 5)) == "35.98765", "value is not correct (rounded to 5 decimal places)"

assert sha1(str(type(sample_float)).encode("utf-8")+b"3f7ac2d1af306a81").hexdigest() == "6728901ff72d3796432f46e0bdba61b4e6bfc918", "type is not float"
assert sha1(str(round(sample_float, 5)).encode("utf-8")+b"3f7ac2d1af306a81").hexdigest() == "4c7fbfc83e04f55c941323502ff1a36d02bf1e63", "value is not correct (rounded to 5 decimal places)"

print('Success!')
```

### list

<!-- grade_id: cell-43678e65bae42fd4, points: 0, grade: false, solution: true -->
```python
# sample_list = [1, 'b', "a", 4.5]

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-c474a81d248cbe37, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_list)) == "<class 'list'>", "type is not list"
assert str(len(sample_list)) == "4", "length is not correct"
assert str(sorted([str(element) for element in sample_list])) == "['1', '4.5', 'a', 'b']", "values are not correct"
assert str(sample_list) == "[1, 'b', 'a', 4.5]", "order of elements is not correct"

assert sha1(str(type(sample_list)).encode("utf-8")+b"9c61b96ffb69f3b5").hexdigest() == "805dd9b064b1c1c5d44d6770cf842845826bf47e", "type is not list"
assert sha1(str(len(sample_list)).encode("utf-8")+b"9c61b96ffb69f3b5").hexdigest() == "f217391d915f55bc13667017a450f1de8f36a31b", "length is not correct"
assert sha1(str(sorted([str(element) for element in sample_list])).encode("utf-8")+b"9c61b96ffb69f3b5").hexdigest() == "5633aff7c4e71311865fd8af5e90fc98ace5f6d1", "values are not correct"
assert sha1(str(sample_list).encode("utf-8")+b"9c61b96ffb69f3b5").hexdigest() == "db26eb952eebb53809f444f0c1d5c9dc2eadcb7a", "order of elements is not correct"

print('Success!')
```

### set

<!-- grade_id: cell-fe5861008f297022, points: 0, grade: false, solution: true -->
```python
# sample_set = {1, 2, 3, 4, 5}

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-5e1c667195b47187, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_set)) == "<class 'set'>", "type is not set"
assert str(len(sample_set)) == "5", "length is not correct"
assert str(sample_set) == "{1, 2, 3, 4, 5}", "value is not correct"

assert sha1(str(type(sample_set)).encode("utf-8")+b"0343c81c3573078f").hexdigest() == "48142b9751028a522c44b4c7fa76819d5c823b9e", "type is not set"
assert sha1(str(len(sample_set)).encode("utf-8")+b"0343c81c3573078f").hexdigest() == "ec1dbbdda921a785b0f86ee71bf3a582b6d3b9e3", "length is not correct"
assert sha1(str(sample_set).encode("utf-8")+b"0343c81c3573078f").hexdigest() == "a7dd246cd6d242937481d659b88fb4bab1536443", "value is not correct"

print('Success!')
```

### tuple

<!-- grade_id: cell-a971551f5f91de4b, points: 0, grade: false, solution: true -->
```python
# sample_tuple = (23, 78)

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-48cfa5ac37cdd98b, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_tuple)) == "<class 'tuple'>", "type is not tuple"
assert str(len(sample_tuple)) == "2", "length is not correct"
assert str(sorted([str(element) for element in sample_tuple])) == "['23', '78']", "values are not correct"
assert str(sample_tuple) == "(23, 78)", "order of elements is not correct"

assert sha1(str(type(sample_tuple)).encode("utf-8")+b"72b47749f1091e41").hexdigest() == "a49b304b06b7b1cf35de2a862d846a1a8859d96b", "type is not tuple"
assert sha1(str(len(sample_tuple)).encode("utf-8")+b"72b47749f1091e41").hexdigest() == "0fdd6ff57b1b1316e3bef8c6e564a33a5b9fa7fe", "length is not correct"
assert sha1(str(sorted([str(element) for element in sample_tuple])).encode("utf-8")+b"72b47749f1091e41").hexdigest() == "e18a7a620b4511fea489918dc572cdd9c9339cd7", "values are not correct"
assert sha1(str(sample_tuple).encode("utf-8")+b"72b47749f1091e41").hexdigest() == "65d6b4c9094ba4f57088c9ac827aba05c0869c0e", "order of elements is not correct"

print('Success!')
```

### dict

<!-- grade_id: cell-ddabbd9234db36c0, points: 0, grade: false, solution: true -->
```python
# sample_dict = {'a': 1, 'b': 2, "c": 3}

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b8b1ad86fbbe578f, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_dict)) == "<class 'dict'>", "type is not dict"
assert str(len(list(sample_dict.keys()))) == "3", "number of keys is not correct"
assert str(sorted(list(sample_dict.keys()))) == "['a', 'b', 'c']", "keys are not correct"
assert str(sorted(list(sample_dict.values()))) == "[1, 2, 3]", "correct keys, but values are not correct"
assert str(sample_dict) == "{'a': 1, 'b': 2, 'c': 3}", "correct keys and values, but incorrect correspondence"

assert sha1(str(type(sample_dict)).encode("utf-8")+b"cd4495e898d547f0").hexdigest() == "5cede4eae2335b275d8fc121b5b6b4fc94bfffc6", "type is not dict"
assert sha1(str(len(list(sample_dict.keys()))).encode("utf-8")+b"cd4495e898d547f0").hexdigest() == "9d7068890cb8ce30a8d65cd46c7e255c232febd2", "number of keys is not correct"
assert sha1(str(sorted(list(sample_dict.keys()))).encode("utf-8")+b"cd4495e898d547f0").hexdigest() == "9fba3acba928b9f5ec4c0fb81d2b16d7cc7d03d9", "keys are not correct"
assert sha1(str(sorted(list(sample_dict.values()))).encode("utf-8")+b"cd4495e898d547f0").hexdigest() == "b0acc66ea4c3b5d8ae4a0fcafb64a656e7e40e09", "correct keys, but values are not correct"
assert sha1(str(sample_dict).encode("utf-8")+b"cd4495e898d547f0").hexdigest() == "f66fd57c7de9563c39543e98bb4713e55657de39", "correct keys and values, but incorrect correspondence"

print('Success!')
```

### bool

<!-- grade_id: cell-e807929d33f8deb6, points: 0, grade: false, solution: true -->
```python
# sample_bool = True

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-cf2f289017808ac8, points: 0, grade: true, solution: false -->
```python
from hashlib import sha1
assert str(type(sample_bool)) == "<class 'bool'>", "type is not bool"
assert str(sample_bool) == "True", "boolean value is not correct"

assert sha1(str(type(sample_bool)).encode("utf-8")+b"317fae7b7b7dcf90").hexdigest() == "9a91d392410c1cb36a4d69728cf5b562b0df8610", "type is not bool"
assert sha1(str(sample_bool).encode("utf-8")+b"317fae7b7b7dcf90").hexdigest() == "f402f680654891eb9a19c875496299f427880132", "boolean value is not correct"

print('Success!')
```

```python

```